--- 
title_meta  : 第一章
title       : 基本介紹
description : "在本章節你初次與R接觸，學習如何使用R做數學運算以及如何宣告變數，也會學習R的基本資料類別，讓我們趕快開始吧！"

--- type:NormalExercise xp:100 skills:1 key:dd8bfba816aff286aa9d14ea7f632bec1e58a1e6
## 熟悉環境

你要在右邊的編輯區輸入程式來完成練習，當你按下Submit Answer的按鈕後，每一行程式都會被執行，R將會回傳正確或錯誤的訊息，程式的輸出結果會呈現在右下角的R Console視窗中。

R就跟Twitter一樣使用 `#` 來加入程式的註解，讓你與其他人可以更快了解這段程式的用途！註解不會被執行故不影響結果，例如在右方編輯區的第一行_# 計算 3 + 4_就是註解。

你也可以直接在R Console視窗中輸入並執行程式，這是一個測試語法是否正確的方式。

*** =instructions
- 在右邊的編輯區已經有了一些範例，你看得出來哪些是會被執行的程式，哪些是不會被執行的註解嗎？
- 加入一行程式計算6加12，然後按下Submit Answer按鈕。

*** =hint
加入一行程式計算6加12，就像範例一樣！

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 計算 3 + 4
3 + 4

# 計算 6 + 12

```

*** =solution
```{r}
# 計算 3 + 4
3 + 4

# 計算 6 + 12
6 + 12
```

*** =sct
```{r}
test_output_contains("7", incorrect_msg = "&#19981;&#38656;&#35201;&#31227;&#38500;&#35336;&#31639;3&#21152;4&#30340;&#31243;&#24335;&#30908;&#65292;&#21152;&#21478;&#22806;&#19968;&#34892;&#31243;&#24335;&#30908;&#35336;&#31639;6&#21152;12&#21363;&#21487;&#12290;")
test_output_contains("18", incorrect_msg = "&#30906;&#23450;&#20320;&#26377;&#21152;&#20102;&#19968;&#34892;&#31243;&#24335;&#30908;&#35336;&#31639;6&#21152;12&#65292;&#27880;&#24847;&#19981;&#35201;&#22312;&#31243;&#24335;&#30908;&#21069;&#38754;&#21152;&#19978; `#`&#65292;&#21542;&#21063;&#23559;&#19981;&#26371;&#34987;&#22519;&#34892;&#65281;")
success_msg("&#22826;&#26834;&#20102;&#65281;&#30475;&#21040;console&#22914;&#20309;&#21576;&#29694;&#20320;&#36865;&#20986;&#30340;&#31243;&#24335;&#20102;&#21966;&#65311;&#29694;&#22312;&#20320;&#24050;&#32147;&#29087;&#24713;&#20102;&#20171;&#38754;&#65292;&#25509;&#19979;&#20358;&#25105;&#20497;&#20358;&#28145;&#20837;&#30740;&#31350;R&#21543;&#65281;")
```

--- type:NormalExercise xp:100 skills:1 key:f50c65ca40332b98e2a74aa17a600d110bc87a81
## R與數學運算

R可以做為一個簡單的計算機，你可以任意使用下列的運算符號：

- 加號： `+`
- 減號： `-`
- 乘號： `*`
- 除號： `/`
- 次方： `^`
- 餘數： `%%`

讓我們多解釋一下最後兩個運算符號：

- `^` 會將其左邊的數字連續乘上右方數字之次數後回傳，舉例來說： `3^2` 運算結果會得到9。
- `%%` 會將左邊的數字除以右邊數字後回傳其餘數，舉例來說： `5 %% 3` 運算結果會得到2。

根據前述的解釋，跟著下列說明完成練習。

*** =instructions
- 在編輯區輸入 `2^5` 來計算2的5次方為多少。
- 在編輯區輸入 `28 %% 6` 來計算28除以6的餘數為多少。
- 點選Submit Answer並在R Console中查看運算結果。
- 注意 `#` 是用來為你的程式加上註解的符號。

*** =hint
另外一個餘數運算範例是： `9 %% 2` ，運算結果會得到 `1`。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 加法
5 + 5 

# 減法
5 - 5 

# 乘法
3 * 5

# 除法
(5 + 5) / 2 

# 次方


# 餘數

```

*** =solution
```{r}
# 加法
5 + 5 

# 減法
5 - 5 

# 乘法
3 * 5

# 除法
(5 + 5) / 2 

# 次方
2 ^ 5

# 餘數
28 %% 6
```

*** =sct
```{r}
msg = "&#19981;&#38656;&#35201;&#31227;&#38500;&#20854;&#20182;&#30340;&#25976;&#23416;&#36939;&#31639;&#31684;&#20363;&#65281;"
test_output_contains("5 + 5", incorrect_msg = msg)
test_output_contains("5 - 5", incorrect_msg = msg)
test_output_contains("(5 + 5)/2", incorrect_msg = msg)
test_output_contains("2^5", incorrect_msg = "&#27425;&#26041;&#30340;&#32244;&#32722;&#19981;&#27491;&#30906;&#65292;&#20877;&#35443;&#32048;&#38321;&#35712;&#19968;&#27425;&#35498;&#26126;&#12290;")
test_output_contains("28 %% 6", incorrect_msg = "&#39192;&#25976;&#30340;&#32244;&#32722;&#19981;&#27491;&#30906;&#65292;&#20877;&#35443;&#32048;&#38321;&#35712;&#19968;&#27425;&#35498;&#26126;&#12290;")
success_msg("&#22826;&#22909;&#20102;&#65281;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#21543;&#12290;")
```


--- type:NormalExercise xp:100 skills:1 key:2ac26323b43f5ce728fb1e29e18d3afc84160fdc
## 宣告變數 

R的基本概念是**變數**。 

變數讓你在R儲存一個值(像是：4) 或者一個物件(像是：函數)，之後你就可以用指派的變數名稱來使用這個值或物件。 

你可以用以下這段程式宣告變數 `my_var` 並且將4這個數值指派給它：

```
my_var <- 4
```

*** =instructions
輪到你囉！在編輯區宣告一個變數 `x` 並且將42這個數值指派給它，完成之後點選Submit Answer。注意當你要求R印出 `x` 時，42就出現了。

*** =hint
觀察一下如何將4指派給 `my_var` 並且在編輯區中做一樣的事情，但是現在將42指派給變數 `x`。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 宣告變數x並將42指派給x
x <- 

# 將變數x的值印出
x
```

*** =solution
```{r}
# 宣告變數x並將42指派給x
x <- 42

# 將變數x的值印出
x
```

*** =sct
```{r}
test_object("x", undefined_msg = "&#30906;&#35469;&#26159;&#21542;&#26377;&#23459;&#21578;&#35722;&#25976; `x`&#12290;",
            incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#26377;&#25351;&#27966;&#27491;&#30906;&#30340;&#25976;&#20540;&#32102; `x`&#12290;") 
success_msg("&#20570;&#24471;&#22909;&#65281;&#20320;&#26159;&#21542;&#26377;&#27880;&#24847;&#21040;&#22312;&#20320;&#23459;&#21578;&#35722;&#25976;&#30340;&#26178;&#20505;&#65292;&#82;&#20006;&#27794;&#26377;&#23559;&#35722;&#25976;&#30340;&#20540;&#21360;&#20986;&#65311;&#32;&#96;&#120;&#32;&#60;&#45;&#32;&#52;&#50;&#96;&#32;&#20006;&#27794;&#26377;&#36664;&#20986;&#65292;&#22240;&#28858;&#82;&#35469;&#28858;&#20320;&#20043;&#24460;&#36996;&#26371;&#20351;&#29992;&#36889;&#20491;&#35722;&#25976;&#65292;&#21542;&#21063;&#20320;&#19981;&#38656;&#35201;&#20808;&#23559;&#36889;&#20491;&#20540;&#25351;&#27966;&#32102;&#35722;&#25976;&#23565;&#21966;&#65311;&#21069;&#36914;&#21040;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:2cb4878a22c17d2f5589992697d96223082a42b9
## 宣告變數 (2)

假設一個水果籃中有5個蘋果，身為一個資料分析師，你希望宣告一個變數 `my_apples` 並將蘋果的個數指派給它。 

*** =instructions
- 在編輯區中輸入： `my_apples <- 5` ，這會將5指派給 `my_apples`。
- 在第二個註解下輸入： `my_apples`，這會在R Console中將 `my_apples` 的值印出。
- 完成之後按Submit Answer按鈕並且觀察R Console：你可以發現5被印出了，所以R已經將 `my_apples` 跟5這個數值連結在一起了。

*** =hint
記得在指派數值或物件給變數時，可以使用 `<-` 來指派，或者使用 `=`，但是R的社群使用者們都偏好 `<-`。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 指派5這個數值給my_apples


# 將my_apples的值印出

```

*** =solution
```{r}
# 指派5這個數值給my_apples
my_apples <- 5

# 將my_apples的值印出
my_apples
```

*** =sct
```{r}
test_object("my_apples", 
            undefined_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#23459;&#21578;&#35722;&#25976; `my_apples`&#12290;",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#25351;&#27966;&#27491;&#30906;&#30340;&#25976;&#20540;&#32102; `my_apples`&#12290;")
test_output_contains("my_apples", 
                     incorrect_msg = "&#20320;&#26159;&#21542;&#26377;&#21578;&#35380;R&#35201;&#23559; `my_apples` &#30340;&#20540;&#21360;&#20986;&#65311;")
success_msg("&#22826;&#22909;&#20102;&#65281;&#21069;&#36914;&#21040;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:a666eb3fb98a7b8cc5930d9c8d2a6997a47561e2
## 宣告變數 (3)

美味的水果籃可不能沒有柳橙！所以你決定加入6個柳橙。身為一個資料分析師，你的直覺反應是宣告變數 `my_oranges` 並將6指派給它。接下來你想計算水果籃裡頭的水果總數，既然你已經完成變數宣告，就可以用清楚的方式來寫程式：
```
my_apples + my_oranges
```

*** =instructions
- 宣告 `my_oranges` 並指派6給它。
- 將 `my_apples` 與 `my_oranges` 相加並將結果印出。
- 要計算水果籃裡頭的水果總數，將 `my_apples` 與 `my_oranges` 相加並指派給一個新變數 `my_fruit`。

*** =hint
`my_fruit` 單純就是 `my_apples` 加 `my_oranges`。只要使用 `+` 就可以進行相加，之後再使用 `<-` 將結果指派給 `my_fruit`。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 宣告my_apples與my_oranges並分別指派蘋果個數為5、柳橙個數為6
my_apples <- 5


# 相加兩個變數然後印出結果


# 宣告變數my_fruit

```

*** =solution
```{r}
# 宣告my_apples與my_oranges並分別指派蘋果個數為5、柳橙個數為6
my_apples  <- 5
my_oranges <- 6

# 相加兩個變數然後印出結果
my_apples + my_oranges

# 宣告變數my_fruit
my_fruit <- my_apples + my_oranges
```

*** =sct
```{r}
test_object("my_apples", 
            undefined_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#23459;&#21578;&#35722;&#25976; `my_apples`&#12290;",
            incorrect_msg = "&#19981;&#35201;&#26356;&#21205;&#35722;&#25976; `my_apples`&#65281;")
test_object("my_oranges", 
            undefined_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#23459;&#21578;&#35722;&#25976; `my_oranges`&#12290;",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#25351;&#27966;&#27491;&#30906;&#30340;&#25976;&#20540;&#32102;&#35722;&#25976; `my_oranges`&#12290;")
test_output_contains("my_apples + my_oranges",
                     incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#26032;&#23531;&#19968;&#27573;&#31243;&#24335; `my_apples + my_oranges`&#65292;&#36889;&#26371;&#23559;&#30456;&#21152;&#32080;&#26524;&#21360;&#20986;&#12290;")
test_object("my_fruit",
            undefined_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#23459;&#21578;&#35722;&#25976;`my_fruit`&#12290;",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#25351;&#27966;&#27491;&#30906;&#30340;&#25976;&#20540;&#32102;&#35722;&#25976; `my_fruit`&#12290;")
success_msg("&#20570;&#24471;&#22826;&#22909;&#20102;&#65281;&#20351;&#29992;&#35722;&#25976;&#20358;&#20570;&#36939;&#31639;&#26368;&#22823;&#30340;&#22909;&#34389;&#23601;&#26159;&#37325;&#35079;&#20351;&#29992;&#65292;&#20551;&#22914;&#20320;&#23559; `my_apples` &#30001;5&#26356;&#25913;&#28858;12&#65292;`my_fruit` &#30340;&#32080;&#26524;&#26371;&#33258;&#21205;&#26356;&#26032;&#12290;&#21069;&#36914;&#21040;&#19979;&#19968;&#20491;&#32244;&#32722;&#21543;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:ff7bb3bd1c38b5db3ce9dde0e45f275163b6a9b6
## 蘋果與柳橙

常識告訴你把蘋果與柳橙相加是很愚蠢的:-) 在前一個練習中， `my_apples` 與 `my_oranges` 這兩個變數各自含有一個數值，在R中 `+` 可以處理數值。假如你指派了一個文字給變數 `my_oranges` (如編輯區所示)，表示你試著將數值與文字相加後指派給變數 `my_fruit`，這不可能行得通。

*** =instructions
- 按Submit Answer按鈕並觀察錯誤訊息，了解為什麼會產生這個錯誤訊息。
- 修正這段程式讓R知道你有6個柳橙，所以水果籃中共有11個水果。

*** =hint
你必須將數值 `6` 指派給變數 `my_oranges` 而非原本的文字 `"six"`。 注意雙引號是用來告訴R `"six"` 是一個文字。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 宣告變數my_apples並將5指派給它
my_apples <- 5 

# 將my_apples的值印出
my_apples       

# 宣告變數my_oranges並將"six"指派給它，然後再把my_oranges印出
my_oranges <- "six" 
my_oranges 

# 宣告變數my_fruit並將my_apples與my_oranges相加後指派給它
my_fruit <- my_apples + my_oranges 
my_fruit
```

*** =solution
```{r}
# 宣告變數my_apples並將5指派給它
my_apples <- 5  

# 將my_apples的值印出
my_apples  

# 宣告變數my_oranges並將"six"指派給它，然後再把my_oranges印出
my_oranges <- 6
my_oranges 

# 宣告變數my_fruit並將my_apples與my_oranges相加後指派給它
my_fruit <- my_apples + my_oranges 
my_fruit
```

*** =sct
```{r}
#test_error()
test_object("my_oranges",
            incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#26377;&#25351;&#27966;&#27491;&#30906;&#30340;&#25976;&#20540;&#32102;&#32;&#96;&#109;&#121;&#95;&#111;&#114;&#97;&#110;&#103;&#101;&#115;&#96;&#12290;&#32;&#23383;&#20018;&#33287;&#25976;&#20540;&#26159;&#19981;&#19968;&#27171;&#30340;&#21908;&#65281;")
test_output_contains("my_apples + my_oranges", 
                     incorrect_msg = "&#36664;&#20986;&#20006;&#27794;&#26377;&#21253;&#21547;&#32;&#96;&#109;&#121;&#95;&#97;&#112;&#112;&#108;&#101;&#115;&#96;&#32;&#21152;&#32;&#96;&#109;&#121;&#95;&#111;&#114;&#97;&#110;&#103;&#101;&#115;&#96;&#30340;&#32080;&#26524;&#12290;")
test_object("my_fruit",
            incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#26377;&#25351;&#27966;&#27491;&#30906;&#30340;&#25976;&#20540;&#32102;&#32;&#96;&#109;&#121;&#95;&#102;&#114;&#117;&#105;&#116;&#96;&#12290;")
success_msg("&#22826;&#26834;&#20102;&#65292;&#35531;&#32380;&#32396;&#20445;&#25345;&#65281;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#12290;")
```


--- type:NormalExercise xp:100 skills:1 key:40c0d484b0ed0826822600b7c1d896d155e3e411
## R的基本資料類別

R可以運算多種的資料類別，其中較為基本的有：

- 有小數點的如 `4.5` 稱為 **數值**。
- 自然數如 `4` 稱為 **整數**，整數同時也屬於數值。
- 布林值(`TRUE` or `FALSE`)稱為 **邏輯值**。
- 字元(或者字串)稱為 **文字**。

用雙引號包起來的就可以讓R知道這是一個文字，像是"一些文字"。

*** =instructions
更改變數的數值：

- 將 `my_numeric` 指派為 `42`。
- 將 `my_character` 指派為 `"universe"`。 用雙引號包起來的就可以讓R知道 `"universe"` 是一個文字。
- 將 `my_logical` 指派為 `FALSE`。

注意英文字母的大小寫在R是有差別的！

*** =hint 
在編輯區將數值重新指派： 
`my_numeric <- 42` 可將42指派給變數 `my_numeric`。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 指派42給my_numeric
my_numeric <- 42.5

# 指派"universe"給my_character
my_character <- "some text"

# 指派FALSE給my_logical
my_logical <- TRUE
```

*** =solution
```{r}
# 指派42給my_numeric
my_numeric <- 42

# 指派"universe"給my_character
my_character <- "universe"

# 指派FALSE給my_logical
my_logical <- FALSE
```

*** =sct
```{r}
test_object("my_numeric", incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#23559;42&#25351;&#27966;&#32102;&#35722;&#25976; `my_numeric`&#65311;")
test_object("my_character", incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#23559; `\"universe\"` &#25351;&#27966;&#32102; `my_character`&#65311;&#21029;&#24536;&#20102;&#29992;&#38617;&#24341;&#34399;&#25226;&#23383;&#20018;&#21253;&#36215;&#20358;&#65281;")
test_object("my_logical", incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#23559; `FALSE` &#25351;&#27966;&#32102; `my_logical`&#65311;&#27880;&#24847; `FALSE` &#27599;&#20491;&#23383;&#27597;&#37117;&#35201;&#20351;&#29992;&#22823;&#23531;&#65281;")
success_msg("&#20570;&#24471;&#22826;&#22909;&#20102;&#65281;&#32380;&#32396;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#21543;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:3c8cc5a936ede8b77b123fda57e937afe5cf1f70
## 如何判斷資料類別？

還記得先前在執行 `5 + "six"` 的時候因為資料類別不符而得到錯誤訊息嗎？其實只要善用右方編輯區的 `class()` 函數事先檢查資料類別就可以避免這種尷尬的情況。

*** =instructions
完成編輯區的程式並將 `my_character` 與 `my_logical` 的資料類別印出。 

*** =hint
參考範例如何將`my_numeric`的資料類別印出，用同樣的方式把 `my_character` 與 `my_logical` 的資料類別也印出來。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 宣告不同種類的變數
my_numeric <- 42
my_character <- "universe"
my_logical <- FALSE 

# 檢查my_numeric的類別
class(my_numeric)

# 檢查my_character的類別


# 檢查my_logical的類別 

```

*** =solution
```{r}
# 宣告不同種類的變數
my_numeric <- 42
my_character <- "universe"
my_logical <- FALSE

# 檢查my_numeric的類別
class(my_numeric)

# 檢查my_character的類別
class(my_character)

# 檢查my_logical的類別
class(my_logical)
```

*** =sct
```{r}
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913;&#35722;&#25976;&#30340;&#23459;&#21578;&#12290;"
lapply(c("my_numeric", "my_character", "my_logical"), test_object, undefined_msg = msg, incorrect_msg = msg)
patt <- "&#30906;&#35469;&#20320;&#26159;&#21542;&#26377;&#20351;&#29992; `class(%1$s)` &#20358;&#21360;&#20986; `%1$s` &#30340;&#36039;&#26009;&#39006;&#21029;&#65311;"
test_output_contains("class(my_numeric)",
                     incorrect_msg = "&#19981;&#38656;&#35201;&#31227;&#38500;&#21360;&#20986; `my_numeric` &#36039;&#26009;&#39006;&#21029;&#30340;&#31243;&#24335;")
test_output_contains("class(my_character)",
                     incorrect_msg = sprintf(patt, "my_character"))
test_output_contains("class(my_logical)",
                     incorrect_msg = sprintf(patt, "my_logical"))
success_msg("&#24685;&#21916;&#20320;&#23436;&#25104;&#20102;&#26412;&#31456;&#31680;&#30340;&#26368;&#24460;&#19968;&#20491;&#32244;&#32722;&#65292;&#35731;&#25105;&#20497;&#21069;&#36914;&#19979;&#20491;&#31456;&#31680;&#20358;&#30740;&#31350;&#21521;&#37327;&#21543;&#65281;")
```



