--- 
title_meta  : 第三章
title       : 矩陣
description : 在本章節你將學會如何在R使用矩陣，包括宣告以及運算，你會用矩陣來分析星際大戰的票房記錄，願原力與你同在！

--- type:NormalExercise xp:100 skills:1 key:61a40bda779e572e5cc17be23357c98bceb73943
## 矩陣是什麼？

矩陣是一種將相同資料類別(數值、文字或邏輯值)排序在列與欄的資料型態，因為有列與欄，所以矩陣具有二個維度。

你可以使用[`matrix()`](http://www.rdocumentation.org/packages/base/functions/matrix)函數創造一個矩陣，例如：

```
matrix(1:9, byrow = TRUE, nrow = 3)
```

在[`matrix()`](http://www.rdocumentation.org/packages/base/functions/matrix)函數中：

- 第一個參數是R會放進列與欄的元素，這裡我們使用 `1:9` 即 `c(1, 2, 3, 4, 5, 6, 7, 8, 9)` 的簡單寫法。
- 參數 `byrow` 表示矩陣在填滿的時候是以列的方向，如果你希望以欄的方向填滿，要調整為 `byrow = FALSE`。 
- 第三個參數 `nrow` 表示矩陣有3列。

*** =instructions
宣告一個具有3列的矩陣並以列的方向將1到9填入。

*** =hint
再詳讀一次說明，答案其實已經很明顯了！

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 宣告一個具有3列的矩陣並將1到9填入
```

*** =solution
```{r}
# 宣告一個具有3列的矩陣並將1到9填入
matrix(1:9, byrow = TRUE, nrow = 3)
```

*** =sct
```{r}
test_function("matrix", c("data", "byrow", "nrow"),
              incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#27491;&#30906;&#22320;&#23459;&#21578;&#30697;&#38499;&#65311;&#20877;&#35443;&#35712;&#19968;&#27425;&#35498;&#26126;&#65292;&#31572;&#26696;&#20854;&#23526;&#24050;&#32147;&#24456;&#26126;&#39023;&#20102;&#65281;")
test_output_contains("matrix(1:9, byrow=TRUE, nrow=3)",
                     incorrect_msg = "&#30475;&#36215;&#20358;&#30697;&#38499;&#30340;&#23459;&#21578;&#26377;&#21839;&#38988;&#65292;&#20877;&#35443;&#35712;&#19968;&#27425;&#35498;&#26126;&#65292;&#31572;&#26696;&#20854;&#23526;&#24050;&#32147;&#24456;&#26126;&#39023;&#20102;&#65281;")
success_msg("&#22909;&#26834;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:befa6552ba5318e95e3c954443aa6ec643d1827f
## 分析矩陣

該是捲起袖子來寫程式的時候了，在接下來的練習中你要分析星際大戰首映票房記錄，願原力與你同在！

有3個向量已經在編輯區被宣告好了，每個向量分別代表著前三部推出的星際大戰電影(曙光乍現、帝國大反擊與絕地大反攻)票房收入，向量中的第一個元素代表美國票房收入，第二個元素代表非美國票房收入(資料來源：維基百科)。

在這個練習中你要合併這三個向量然後建立一個矩陣。

*** =instructions
- 利用 `c(new_hope, empire_strikes, return_jedi)` 合併這三個向量並指派給一個新向量 `box_office`。
- 使用 `matrix()` 函數建立一個擁有3列的矩陣，每一列都代表一部電影，第一個參數為 `box_office` 它包含所有票房收入記錄，第二個與第三個參數你要指定 `nrow = 3` 以及 `byrow = TRUE`。命名這個新建立的矩陣為 `star_wars_matrix`。

*** =hint
- `box_office <- c(new_hope, empire_strikes, return_jedi)` 可以合併成一個有6個元素的向量。
- `matrix(box_office, nrow = ..., by_row ...)` 是建立矩陣的函數架構。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 星際大戰票房(百萬元美金！)
new_hope <- c(460.998, 314.4)
empire_strikes <- c(290.475, 247.900)
return_jedi <- c(309.306, 165.8)

# 宣告box_office
box_office <- 

# 建立star_wars_matrix
star_wars_matrix <- 
```

*** =solution
```{r}
# 星際大戰票房(百萬元美金！)
new_hope <- c(460.998, 314.4)
empire_strikes <- c(290.475, 247.900)
return_jedi <- c(309.306, 165.8)

# 宣告box_office
box_office <- c(new_hope, empire_strikes, return_jedi)

# 建立star_wars_matrix
star_wars_matrix <- matrix(box_office, nrow = 3, byrow = TRUE) 
```

*** =sct
```{r}
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913; `new_hope` &#12289; `empire_strikes` &#33287; `return_jedi`&#12290;"
test_object("new_hope", undefined_msg = msg, incorrect_msg = msg)
test_object("empire_strikes", undefined_msg = msg, incorrect_msg = msg)
test_object("return_jedi", undefined_msg = msg, incorrect_msg = msg)

test_object("box_office", incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#27491;&#30906;&#21512;&#20341; `new_hope` &#12289; `empire_strikes` &#33287; `return_jedi` &#25104;&#28858; `box_office`&#65311;")

test_function("matrix", c("data", "nrow", "byrow"),
              incorrect_msg = "&#35531;&#30906;&#35469;&#12288;`matrix()`&#12288;&#30340;&#21443;&#25976;&#65306; `box_office`&#12288;&#12289; `nrow = 3`&#12288;&#12289; `by_row = TRUE`&#12290;")

test_object("star_wars_matrix",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#25351;&#27966; `matrix()` &#30340;&#32080;&#26524;&#32102; `star_wars_matrix`&#65311;")

success_msg("&#21407;&#21147;&#30495;&#30340;&#33287;&#20320;&#21516;&#22312;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:37aed0b157460a60af201e3dc57142bbbf187c71
## 為矩陣命名

為了記住 `star_wars_matrix` 中儲存的資料，你打算為列與欄命名。

```
rownames(my_matrix) <- row_names_vector
colnames(my_matrix) <- col_names_vector
```

我們準備了2個向量 `region` 與 `titles` 。你會需要用這兩個向量來分別為 `star_wars_matrix` 的列與欄命名。

*** =instructions
- 使用 `colnames()` 將 `region` 指派給 `star_wars_matrix` 的欄名稱。
- 使用 `rownames()` 將 `titles` 指派給 `star_wars_matrix` 的列名稱。
- 印出 `star_wars_matrix` 看結果。

*** =hint
你可以使用 `colnames(star_wars_matrix) <- region` 來為 `star_wars_matrix` 命名，列的命名也很類似。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 星際大戰票房(百萬元美金！)
new_hope <- c(460.998, 314.4)
empire_strikes <- c(290.475, 247.900)
return_jedi <- c(309.306, 165.8)

# 建立star_wars_matrix
star_wars_matrix <- matrix(c(new_hope, empire_strikes, return_jedi), nrow = 3, byrow = TRUE)

# 用來命名的向量region與titles
region <- c("US", "non-US")
titles <- c("A New Hope", "The Empire Strikes Back", "Return of the Jedi")

# 用region為欄命名 


# 用titles為列命名


# 印出star_wars_matrix
```

*** =solution
```{r}
# 星際大戰票房(百萬元美金！)
new_hope <- c(460.998, 314.4)
empire_strikes <- c(290.475, 247.900)
return_jedi <- c(309.306, 165.8)

# 建立star_wars_matrix
star_wars_matrix <- matrix(c(new_hope, empire_strikes, return_jedi), nrow = 3, byrow = TRUE)

# 用來命名的向量region與titles
region <- c("US", "non-US")
titles <- c("A New Hope", "The Empire Strikes Back", "Return of the Jedi")

# 用region為欄命名
colnames(star_wars_matrix) <- region

# 用titles為列命名
rownames(star_wars_matrix) <- titles

# 印出star_wars_matrix
star_wars_matrix
```

*** =sct
```{r}
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913; `new_hope` &#12289; `empire_strikes` &#33287; `return_jedi`&#12290;"
test_object("new_hope", undefined_msg = msg, incorrect_msg = msg)
test_object("empire_strikes", undefined_msg = msg, incorrect_msg = msg)
test_object("return_jedi", undefined_msg = msg, incorrect_msg = msg)
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913; `star_wars_matrix` &#65292;&#21482;&#35201;&#24171;&#21015;&#33287;&#27396;&#21629;&#21517;&#12290;" 
test_object("star_wars_matrix", incorrect_msg = msg)
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913; `region` &#33287; `titles` &#36889;2&#20491;&#24050;&#32147;&#24171;&#20320;&#23450;&#32681;&#22909;&#30340;&#21521;&#37327;&#12290;"
test_object("region", undefined_msg = msg, incorrect_msg = msg)
test_object("titles", undefined_msg = msg, incorrect_msg = msg)
test_object("star_wars_matrix", eq_condition = "equal",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#26377;&#27491;&#30906;&#21629;&#21517; `star_wars_matrix` &#30340;&#21015;&#33287;&#27396;&#65311; &#20351;&#29992; `colnames(star_wars_matrix) <- region` &#28858;&#27396;&#21629;&#21517;&#65292;&#29992;&#21516;&#27171;&#26041;&#27861;&#23565;&#21015;&#21629;&#21517;&#12290;")
test_output_contains("star_wars_matrix", incorrect_msg = "&#21029;&#24536;&#20102;&#35201;&#25226; `star_wars_matrix` &#21360;&#20986;&#20358;&#30475;&#30475;&#12290;")
success_msg("&#22826;&#26834;&#20102;&#65292;&#20320;&#27491;&#22312;&#25104;&#28858;&#19968;&#20491;R&#32085;&#22320;&#27494;&#22763;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:99052c568114e85002519d338103ed0c8521ce61
## 計算全球票房

全球票房收入是在好萊塢成為傳奇電影的首要條件。

你必須知道如何把美國票房收入與非美國票房收入相加才能計算這三部星際大戰電影的全球票房收入。

R的[`rowSums()`](http://www.rdocumentation.org/packages/base/functions/colSums)函數可以方便你計算矩陣中每一列的總和，它會產生一個新的向量：

```
rowSums(my_matrix)
```

*** =instructions
計算三部電影的全球票房收入並指派給向量 `worldwide_vector`。

*** =hint
`rowSums(star_wars_matrix)` 可以計算矩陣每一列的總和，即每一部電影的全球票房收入。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 建立star_wars_matrix
box_office <- c(460.998, 314.4, 290.475, 247.900, 309.306, 165.8)
star_wars_matrix <- matrix(box_office, nrow = 3, byrow = TRUE,
                           dimnames = list(c("A New Hope", "The Empire Strikes Back", "Return of the Jedi"), 
                                           c("US", "non-US")))

# 計算全球票房收入
worldwide_vector <- 
```

*** =solution
```{r}
# 建立star_wars_matrix
box_office <- c(460.998, 314.4, 290.475, 247.900, 309.306, 165.8)
star_wars_matrix <- matrix(box_office, nrow = 3, byrow = TRUE,
                           dimnames = list(c("A New Hope", "The Empire Strikes Back", "Return of the Jedi"), 
                                           c("US", "non-US")))

# 計算全球票房收入
worldwide_vector <- rowSums(star_wars_matrix)
```

*** =sct
```{r}
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913; `box_office_all` &#33287; `star_wars_marix`&#12290;"
test_object("box_office", undefined_msg = msg, incorrect_msg = msg)
test_object("star_wars_matrix", undefined_msg = msg, incorrect_msg = msg)
test_object("worldwide_vector", incorrect_msg = "&#23565; `star_wars_matrix` &#20351;&#29992; `rowSums()` &#20989;&#25976;&#20006;&#23559;&#32080;&#26524;&#25351;&#27966;&#32102; `worldwide_vector`&#12290;")
success_msg("&#20570;&#24471;&#22909;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:b3920821c2964147c5db451c4d66c8e6172dfda4
## 再多加一欄 

前一個練習中你計算了3部星際大戰電影的全球票房收入，但是這個向量沒有合併至 `star_wars_matrix` 。

使用[`cbind()`](http://www.rdocumentation.org/packages/base/functions/cbind)函數將一個或更多的欄加入矩陣，例如：

```
big_matrix <- cbind(matrix1, matrix2, vector1 ...)
```

*** =instructions
使用[`cbind()`](http://www.rdocumentation.org/packages/base/functions/cbind)函數加入 `worldwide_vector` 做為 `star_wars_matrix` 中新的一欄，並將結果指派給 `all_wars_matrix`。

*** =hint
將2個變數 `star_wars_matrix` 與 `worldwide_vector` 做為 `cbind()` 的參數然後指派給 `all_wars_matrix`。

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 建立star_wars_matrix
box_office <- c(460.998, 314.4, 290.475, 247.900, 309.306, 165.8)
star_wars_matrix <- matrix(box_office, nrow = 3, byrow = TRUE,
                           dimnames = list(c("A New Hope", "The Empire Strikes Back", "Return of the Jedi"), 
                                           c("US", "non-US")))

# 計算全球票房收入
worldwide_vector <- rowSums(star_wars_matrix)

# 將worldwide_vector加入star_wars_matrix
all_wars_matrix <- 
```

*** =solution
```{r}
# 建立star_wars_matrix
box_office <- c(460.998, 314.4, 290.475, 247.900, 309.306, 165.8)
star_wars_matrix <- matrix(box_office, nrow = 3, byrow = TRUE,
                           dimnames = list(c("A New Hope", "The Empire Strikes Back", "Return of the Jedi"), 
                                           c("US", "non-US")))

# 計算全球票房收入
worldwide_vector <- rowSums(star_wars_matrix)

# 將worldwide_vector加入star_wars_matrix
all_wars_matrix <- cbind(star_wars_matrix, worldwide_vector)
```

*** =sct
```{r}
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913; `box_office_all` &#33287; `star_wars_marix`&#12290;"
test_object("box_office", undefined_msg = msg, incorrect_msg = msg)
test_object("star_wars_matrix", undefined_msg = msg, incorrect_msg = msg)
test_object("worldwide_vector",
            incorrect_msg = "&#23559; `rowSums(star_wars_matrix)` &#30340;&#32080;&#26524;&#20786;&#23384;&#22312; `worldwide_vector` &#20013;&#12290;")

# TODO: Fix in testwhat
msg <- "&#30906;&#35469;&#26159;&#21542;&#27491;&#30906;&#20351;&#29992; `cbind()` &#20358;&#21512;&#20341; `worldwide_vector` &#33287; `star_wars_matrix`&#65311;&#20320;&#25033;&#35442;&#20381;&#24207;&#23559; `star_wars_matrix` &#33287; `world_wide_vector` &#20570;&#28858; `cbind()` &#30340;&#20841;&#20491;&#21443;&#25976;&#65292;&#22914;&#27492;&#19968;&#20358; `all_wars_matrix` &#23559;&#26377;&#19977;&#21015;&#33287;&#19977;&#27396;&#12290;"
test_object("all_wars_matrix", eval = FALSE)
test_what(
  expect_equal(
    dim(get("all_wars_matrix", get_solution_env())), 
    dim(get("all_wars_matrix", .GlobalEnv))),
  feedback = msg)

test_object("all_wars_matrix", incorrect_msg = msg)
success_msg("&#20570;&#24471;&#22909;&#65292;&#25509;&#19979;&#20358;&#25105;&#20497;&#20358;&#35430;&#33879;&#21152;&#21015;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;");
```


--- type:NormalExercise xp:100 skills:1 key:31368f587442b9db1b8836534e6998b0d024835e
## 新增一列

就如同作用力與反作用力，有[`cbind()`](http://www.rdocumentation.org/packages/base/functions/cbind)自然也會有[`rbind()`](http://www.rdocumentation.org/packages/base/functions/cbind)(我們實在不太會用譬喻法。) 

在R工作區已經宣告好兩個矩陣([工作區是什麼](http://www.statmethods.net/interface/workspace.html))：

- 我們一直在使用的第一個三部曲電影矩陣 `star_wars_matrix`。
- 第二個三部曲電影矩陣`star_wars_matrix2`。 

在R Console輸入 `ls()` 可以檢視工作區的內容，輸入已經宣告好的2個矩陣可以看到內容。

*** =instructions
用 `rbind()` 把 `star_wars_matrix` 和 `star_wars_matrix2` 合併並將結果指派給 `all_wars_matrix`。

*** =hint
合併兩個矩陣：
```
rbind(matrix1, matrix2)
```
把結果指派給 `all_wars_matrix`。


*** =pre_exercise_code
```{r}
# 建立star_wars_matrix
box_office_all <- c(461, 314.4, 290.5, 247.9, 309.3, 165.8)
movie_names <- c("A New Hope","The Empire Strikes Back","Return of the Jedi")
col_titles <- c("US","non-US")
star_wars_matrix <- matrix(box_office_all, nrow = 3, byrow = TRUE, dimnames = list(movie_names, col_titles))

# 建立star_wars_matrix2
box_office_all2 <- c(474.5, 552.5, 310.7, 338.7, 380.3, 468.5)
movie_names2 <- c("The Phantom Menace", "Attack of the Clones", "Revenge of the Sith")
star_wars_matrix2 <- matrix(box_office_all2, nrow=3, byrow = TRUE, dimnames = list(movie_names2, col_titles))

# 移除所有物件只留下all_wars_matrix
rm(box_office_all)
rm(movie_names)
rm(col_titles)
rm(box_office_all2)
rm(movie_names2)
```

*** =sample_code
```{r}
# star_wars_matrix與star_wars_matrix2已經存在於工作區中
star_wars_matrix  
star_wars_matrix2 

# 合併2個三部曲
all_wars_matrix <- 
```

*** =solution
```{r}
# star_wars_matrix與star_wars_matrix2已經存在於工作區中
star_wars_matrix  
star_wars_matrix2 

# 合併2個三部曲
all_wars_matrix <- rbind(star_wars_matrix, star_wars_matrix2)
```

*** =sct
```{r}
msg = "&#19981;&#38656;&#35201;&#26356;&#25913;`star_wars_matrix` &#33287; `star_wars_matrix2`&#12290;"
test_object("star_wars_matrix", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("star_wars_matrix2", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)

# Fix until testwhat gets updated again
# TODO: Fix in testwhat
test_object("all_wars_matrix", eval = FALSE)
test_what(
  expect_equal(
    dim(get("all_wars_matrix", get_solution_env())), 
    dim(get("all_wars_matrix", .GlobalEnv))),
  feedback = 
    paste("&#35531;&#30906;&#35469; `rbind()` &#21443;&#25976;&#20013;&#30340;2&#20491;&#21443;&#25976;&#65306; `star_wars_matrix`",
          "&#33287; `star_wars_matrix2`&#65292;&#32080;&#26524;&#25033;&#35442;&#26159;&#19968;&#20491;&#20855;&#26377;",
          "6&#21015;&#33287;2&#27396;&#30340;&#30697;&#38499;&#12290;"))

test_object("all_wars_matrix", incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#27491;&#30906;&#20351;&#29992; `rbind()` &#21109;&#36896; `all_wars_matrix()`&#65311; `rbind()` &#25033;&#35442;&#35201;&#20381;&#24207;&#26377;2&#20491;&#21443;&#25976;&#65306; `star_wars_matrix` &#33287; `star_wars_matrix2`&#12290;")

success_msg("&#22826;&#23436;&#32654;&#20102;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#20006;&#21152;&#20837; `colSums()` &#20989;&#25976;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:7dfd6a4e59e15742efd72fee93321ec289165846
## 星際大戰六部曲票房總收入

如同[`cbind()`](http://www.rdocumentation.org/packages/base/functions/cbind)與[`rbind()`](http://www.rdocumentation.org/packages/base/functions/cbind)，[`colSums()`](http://www.rdocumentation.org/packages/base/functions/colSums)也有對應的[`rowSums()`](http://www.rdocumentation.org/packages/base/functions/colSums)。R工作區已經有你在前一個練習建立的 `all_wars_matrix`， 輸入 `all_wars_matrix` 看看，接著我們要計算六部曲的票房總收入。

*** =instructions
- 計算美國與非美國的票房總收入並指派給 `total_revenue_vector`，你可以使用[`colSums()`](http://www.rdocumentation.org/packages/base/functions/colSums)函數。
- 印出 `total_revenue_vector` 看看結果。

*** =hint
使用[`colSums()`](http://www.rdocumentation.org/packages/base/functions/colSums)函數function，以 `star_wars_matrix` 做為參數找出各區域票房收入。

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/all_wars_matrix.RData"))
```

*** =sample_code
```{r}
# all_wars_matrix已經存在於工作區中
all_wars_matrix

# 美國與非美國的票房收入
total_revenue_vector <- 
  
# 印出total_revenue_vector
```

*** =solution
```{r}
# all_wars_matrix已經存在於工作區中
all_wars_matrix

# 美國與非美國的票房收入
total_revenue_vector <- colSums(all_wars_matrix)

# 印出total_revenue_vector
total_revenue_vector
```

*** =sct
```{r}
msg = "&#19981;&#38656;&#35201;&#26356;&#25913;&#24037;&#20316;&#21312;&#24050;&#24314;&#31435;&#22909;&#30340; `all_wars_matrix`&#12290;"
test_object("all_wars_matrix", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_function("colSums", "x", incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#23565;all_wars_matrix&#20351;&#29992; `colSums()` &#20989;&#25976;&#65311;")
test_object("total_revenue_vector",
            incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#23559; `colSums(all_wars_matrix)` &#25351;&#27966;&#32102; `total_revenue_vector`&#65311;")
test_output_contains("total_revenue_vector", incorrect_msg = "&#21029;&#24536;&#35352;&#25226; `total_revenue_vector` &#21360;&#20986;&#20358;&#65281;")
success_msg("&#22826;&#32654;&#20102;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#30697;&#38499;&#30340;&#36984;&#25799;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:3f713b0c9ff675bf457941041819813ad7a5c19d
## 選擇矩陣中的元素

跟向量一樣都是使用中括號 `[ ]` 來選擇矩陣中的一個或多個元素，因為矩陣具有2個維度，你應該使用逗號分格哪些列要選、哪些欄要選，例如：

- `my_matrix[1,2]` 可以將矩陣中第一列與第二欄的元素選出。
- `my_matrix[1:3,2:4]` 可以將矩陣中第一列至第三列與第二欄至第四欄的資料選出。

如果你希望選出列或欄的所有元素，就不需要在逗號前後加上數字，例如：

- `my_matrix[,1]` 可以把第一欄所有元素選出。
- `my_matrix[1,]` 可以把第一列所有元素選出。

讓我們繼續回到星際大戰，在前一個練習中 `all_wars_matrix` 已經存在於工作區中。

*** =instructions
- 把六部曲非美國的電影票房收入( `all_wars_matrix` 中第二欄所有的元素)指派給 `non_us_all`。
- 對 `non_us_all` 使用 `mean()` 函數來計算非美國票房收入的平均值並簡單印出結果。
- 從 `all_wars_matrix` 選出前兩部電影的非美國票房收入並指派給 `non_us_some`。
- 對 `non_us_some` 使用 `mean()` 函數計算平均值並印出。

*** =hint
用 `my_matrix[,2]` 可以將 `my_matrix` 的第二欄中所有元素選出。

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/all_wars_matrix.RData"))
```

*** =sample_code
```{r}
# all_wars_matrix已經存在於工作區中
all_wars_matrix

# 選出全部電影的非美國票房收入
non_us_all <- 
  
# 平均非美國票房收入

  
# 選出前兩部電影的非美國票房收入
non_us_some <- 
  
# 前兩部電影的非美國票房平均收入

```

*** =solution
```{r}
# all_wars_matrix已經存在於工作區中
all_wars_matrix

# 選出全部電影的非美國票房收入
non_us_all <- all_wars_matrix[,2]
  
# 平均非美國票房收入
mean(non_us_all)
  
# 選出前兩部電影的非美國票房收入
non_us_some <- all_wars_matrix[1:2,2]
  
# 前兩部電影的非美國票房平均收入
mean(non_us_some)
```

*** =sct
```{r}
msg = "&#19981;&#38656;&#35201;&#26356;&#25913;&#24037;&#20316;&#21312;&#24050;&#24314;&#31435;&#22909;&#30340; `all_wars_matrix`&#12290;"
test_object("all_wars_matrix", undefined_msg = msg, incorrect_msg = msg)

test_object("non_us_all",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#23559; `all_wars_matrix` &#30340;&#31532;&#20108;&#27396;&#25351;&#27966;&#32102; `non_us_all` &#65311;&#20320;&#21487;&#20197;&#20351;&#29992; `[, 2]` &#36774;&#21040;&#65281;")
test_output_contains("mean(non_us_all)",
                     incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#20351;&#29992; `mean(non_us_all)` &#35336;&#31639; `non_us_all` &#30340;&#24179;&#22343;&#20540;&#20006;&#21360;&#20986;&#32080;&#26524;&#65311;")
test_object("non_us_some",
            incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#25351;&#27966;&#21069;2&#37096;&#38651;&#24433;&#30340;&#38750;&#32654;&#22283;&#31080;&#25151;&#25910;&#20837;&#32102; `non_us_some` &#65311;&#20320;&#21487;&#20197;&#20351;&#29992; `[1:2,2]` &#36774;&#21040;&#65281;")
test_output_contains("mean(non_us_some)",
                     incorrect_msg = "&#35531;&#30906;&#35469;&#26159;&#21542;&#20351;&#29992; `mean(non_us_some)` &#35336;&#31639; `non_us_some` &#30340;&#24179;&#22343;&#20540;&#20006;&#21360;&#20986;&#32080;&#26524;&#65311;")
success_msg("&#22826;&#21426;&#23475;&#20102;&#65292;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:4859ff893aeb0e78190fbece8655b0a65b58a814
## 一點點矩陣數學運算

跟你在向量學到的相同，標準運算子 `+` 、 `-` 、 `/` 、 `*` 等等都可以用在矩陣計算。

例如 `2 * my_matrix` 會將每個 `my_matrix` 中的元素均乘以2。

身為盧卡斯影業新聘雇的資料分析師，你想找出不同地理區的六部曲觀看人數。你已經建立好總票房收入的 `all_wars_matrix`，現在假設電影票價格為5元，將票房收入除以5就可以得到觀看人數。

*** =instructions
- 將 `all_wars_matrix` 除以5得到觀看人數(百萬人)，指派結果給 `visitors`。
- 印出 `visitors` 看看。

*** =hint
觀看人數為 `all_wars_matrix` 除以5。

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/all_wars_matrix.RData"))
``` 

*** =sample_code
```{r}
# all_wars_matrix已經存在於工作區中
all_wars_matrix

# 估計觀看人數
visitors <- 
  
# 印出觀看人數

```

*** =solution
```{r}
# all_wars_matrix已經存在於工作區中
all_wars_matrix

# 估計觀看人數
visitors <- all_wars_matrix / 5

# 印出觀看人數
visitors
```

*** =sct
```{r}
msg = "&#19981;&#38656;&#35201;&#26356;&#25913;&#24037;&#20316;&#21312;&#24050;&#24314;&#31435;&#22909;&#30340; `all_wars_matrix`&#12290;"
test_object("all_wars_matrix", undefined_msg = msg, incorrect_msg = msg)

test_object("visitors",
            incorrect_msg = "&#30475;&#36215;&#20358; `visitors` &#19981;&#27491;&#30906;&#65292;&#23559; `all_wars_matrix` &#38500;&#20197;5&#20006;&#23559;&#32080;&#26524;&#25351;&#27966;&#32102; `visitors`&#12290;")
test_output_contains("visitors", incorrect_msg = "&#21029;&#24536;&#20102;&#23559; `visitors` &#21360;&#20986;&#20358;&#30475;&#30475;&#12290;")
success_msg("&#35336;&#31639;&#32080;&#26524;&#21578;&#35380;&#20320;&#65292;&#26143;&#38555;&#22823;&#25136;&#22235;&#37096;&#26354;&#65306;&#26329;&#20809;&#20045;&#29694;&#20809;&#26159;&#22312;&#32654;&#22283;&#23601;&#26377;9&#21315;2&#30334;&#33836;&#20491;&#35264;&#30475;&#20154;&#25976;&#65292;&#24456;&#39514;&#20154;&#23565;&#21966;&#65311;&#21069;&#36914;&#19979;&#19968;&#20491;&#32244;&#32722;&#65281;")
```


--- type:NormalExercise xp:100 skills:1 key:6c34c3c66d9f2d2996cb8288112f2e604feac8ab
## 一點點矩陣數學運算 (2) 

如同 `2 * my_matrix` 會將 `my_matrix` 中的每個元素都乘以2， `my_matrix1 * my_matrix2` 會產生一個對應元素相乘的矩陣。 

老闆盧卡斯看了前一個練習的結果後指出電影票價其實是有隨著時間調漲，要求你利用 `ticket_prices_matrix` 中的票價資料(資料來源：imagination) 重做一次分析。

_對於矩陣熟悉的使用者要注意這裡不是使用標準矩陣運算子，標準矩陣運算子要使用 `%*%`。_

*** =instructions
- `all_wars_matrix` 除以 `ticket_prices_matrix` 可以得到美國與非美國的六部曲觀看人數，將結果指派給 `visitors`。
- 選擇 `visitors` 中代表美國觀看人數的第一欄資料指派給 `us_visitors`。
- 計算美國的平均觀看人數並印出結果。

*** =hint
- 使用 `mean()` 函數計算。
- 使用 `visitors[ ,1]` 來取出 `visitors` 的第一欄以得到美國的觀看人數。

*** =pre_exercise_code
```{r}
load(url("http://s3.amazonaws.com/assets.datacamp.com/course/intro_to_r/all_wars_matrix.RData"))
movie_names <- c("A New Hope","The Empire Strikes Back","Return of the Jedi", "The Phantom Menace", "Attack of the Clones", "Revenge of the Sith")
col_titles <- c("US","non-US")
ticket_prices_matrix <- matrix(c(5, 5, 6, 6, 7, 7, 4, 4, 4.5, 4.5, 4.9, 4.9), nrow = 6, byrow = TRUE, dimnames = list(movie_names,col_titles))
```

*** =sample_code
```{r}
# all_wars_matrix與ticket_prices_matrix已建立於工作區
all_wars_matrix
ticket_prices_matrix

# 估計觀看人數
visitors <- 

# 美國觀看人數
us_visitors <- 

# 美國平均觀看人數

```

*** =solution
```{r}
# all_wars_matrix與ticket_prices_matrix已建立於工作區
all_wars_matrix
ticket_prices_matrix

# 估計觀看人數
visitors <- all_wars_matrix / ticket_prices_matrix

# 美國觀看人數
us_visitors <- visitors[ ,1]

# 美國平均觀看人數
mean(us_visitors)
```

*** =sct
```{r}
msg <- "&#19981;&#38656;&#35201;&#26356;&#25913;&#24037;&#20316;&#21312;&#24050;&#24314;&#31435;&#22909;&#30340; `all_wars_matrix`&#12290;"
test_object("all_wars_matrix", undefined_msg = msg, incorrect_msg = msg)
test_object("ticket_prices_matrix", undefined_msg = msg, incorrect_msg = msg)

test_object("visitors",
            incorrect_msg = "&#30906;&#35469;&#26159;&#21542;&#27491;&#30906;&#24314;&#31435;&#30697;&#38499; `visitors`&#65311; &#23559; `all_wars_matrix` &#38500;&#20197; `ticket_prices_matrix` &#23601;&#21487;&#20197;&#31639;&#20986;&#12290;")
test_object("us_visitors", incorrect_msg = "&#20320;&#21487;&#20197;&#20351;&#29992; `[,1]` &#20358;&#36984;&#20986; `visitors` &#30340;&#31532;&#19968;&#27396;&#20006;&#25351;&#27966;&#32102; `us_visitors`&#65281;")
test_output_contains("mean(us_visitors)", incorrect_msg = "&#19968;&#26086;&#20320;&#24314;&#31435;&#22909; `us_visitors` &#23601;&#21487;&#20197;&#20351;&#29992; `mean()` &#20358;&#35336;&#31639;&#32654;&#22283;&#24179;&#22343;&#35264;&#30475;&#20154;&#25976;&#20006;&#35352;&#24471;&#35201;&#25226;&#32080;&#26524;&#21360;&#20986;&#20358;&#12290;")

success_msg("R&#30340;&#21407;&#21147;&#33287;&#20320;&#21516;&#22312;&#65292;&#36889;&#26159;&#30697;&#38499;&#21934;&#20803;&#30340;&#26368;&#24460;&#19968;&#20491;&#32244;&#32722;&#65292;&#19979;&#19968;&#20491;&#31456;&#31680;&#26159;&#22240;&#32032;&#21521;&#37327;&#65281;")
```


